```python
# Initialize Otter
import otter
grader = otter.Notebook("lab4.ipynb")
```

# Coding lab \#4

```{grid} 2 2 2 2 
**2D elements** 
:::{figure} figures/AshesiETHLogo.png
:align: right
:width: 300px
:::
```

:::{admonition} How to use this notebook?
This notebook is made of text cells and code cells. The code cells have to be **executed** to see the result of the program. \
To execute a cell, simply select it and click on the "play" button \(&#9658;\) in the tool bar just above the notebook, or type `shift+enter`
It is important to execute the code cells in their order of appearance in the notebook.

**As a first step**: execute the very first OtterGrader code cell at the very top of this notebook. 

If you are running into any problems, try restarting the kernel and then see if the issue persist. 
:::

## Learning Goals

After using this notebook, you should be able to:
* perform FE computations under the plane stress assumptions using CST (Constant Strain Triangular) elements
* understand why mesh refinement is necessary

## Problem Statement

The objective of this coding lab is to write your own finite element code to simulate linear elastic boundary value problems in 2D (under plane stress assumptions). 
Plane stress means that we are dealing with a thin object loaded within its own median plane. 
Examples include bi-axial testing of thin samples, or a spanner used to tighten a bolt.
To  model these structures, we will use constant strain triangles (CSTs), whose formulation was derived in Section 16 of the lecture notes (see especially Sections 16.3 and 16.3.1 for the $\boldsymbol{B}$-matrix, stiffness matrix, internal forces, etc. and Section 1.4 (particularly equations 1.32) for matrix $\boldsymbol{E}=\mathbb{C}$).

%<img src="figures/P4_spanner04.png" alt="spanner" style="width:200px;"/>
%<!-- <p style="font-style:italic;">Figure 1: spanner </p> -->

```{figure} figures/P4_spanner04.png
:name: spannerFE
:alt: spanner FE
:class: bg-primary
:width: 100px"
:align: center

```

<div style="text-align: center"> <sup><sub> image credit: http://mae.uta.edu/~lawrence/ansys/solutions/wrench/wrench.htm</sub></sup>  </div>

___
## problem 1 a)


:::{admonition} submission requirement
:class: tip

A person is trying to balance on a road sign (a very bad idea!), which is fully clamped on the left-hand-side (along $x_1 = 0$).
At a point in time the person's weight distribution can be approximated by a distributed load as shown in the figure below.
To model the deformation of the road sign under this distributed load with the finite element method, we decide to discretize the domain into 16 CSTs, as shown in the figure below. 

```{figure} figures/lab4_plate_nodes_0.png
:name: roadSign
:alt: plane stress road sign
:class: bg-primary
:width: 400px
:align: center

```
We now solve for the deformation of and stresses within the road sign under the load by the FEM (assuming plane stress). 
The dimensions are given by $a = 0.25$m and $t=0.1$m, Young's modulus and Poisson's ratio are $E_y = 2$GPa and $\nu = 0.3$, respectively.
*Assume* that the external forces on nodes 1, 2 and 3 (equivalent to the distributed load) have components 

$F^1_2 = -1000$N, $F^2_2 = -2000$N, and $F^3_2 = -800$N. 

Follow the intstructions in the coding section below and input all missing information as indicated, to solve this problem under the **plane stress** assumption using **CST**s. Then follow the instructions in the pre-processing section to complete problem 1. 
:::

___

# coding

For this, we need to import some useful Python libraries.

**Execute the code cell below** so that the necessary libraries are imported.

```python
import numpy as np
import matplotlib.pyplot as plt
from matplotlib.patches import Polygon
from matplotlib.collections import PatchCollection
from matplotlib.legend_handler import HandlerTuple
import os
import subprocess
import meshio
from lib.plotting_functions import plot_original_and_displaced_structure, plot_von_Mises_stress_distribution,plot_h_vs_Sigma
from scipy.sparse import csr_matrix
from scipy.sparse.linalg import spsolve
# Enable inline plotting in Jupyter Notebook
%matplotlib inline
# And display a message once all libraries are imported
print("Libraries imported.")
```

Output:
```
Libraries imported.
```

### pre-processing

```python
def input_definition(problem):
    ###########################################################################################################
    if problem == 1: 

        # fill out all the missing parts:
        a = 0.25
        # coordinates of the global nodes:
        nodal_positions = np.array([[0,2*a],
                                    [a, 2*a],
                                    [2*a, 2*a],
                                    [3*a, 2*a],
                                    [4*a, 2*a],
                                    [0, a],
                                    [a, a],
                                    [2*a, a],
                                    [3*a, a],
                                    [4*a, a],
                                    [0, 0],
                                    [a, 0],
                                    [2*a, 0],
                                    [3*a, 0],
                                    [4*a, 0]])
        
        # local-to-global map / node connectivity
        nodal_connectivities = np.array([[0, 5, 1], 
                                         [1, 5, 6],
                                         [1, 6, 2],
                                         [2, 6, 7],
                                         [2, 7, 3],
                                         [3, 7, 8],
                                         [3, 8, 4],
                                         [4, 8, 9],
                                         [5, 10, 6],
                                         [6, 10, 11],
                                         [6, 11, 7],
                                         [7, 11, 12],
                                         [7, 12, 8],
                                         [8, 12, 13],
                                         [8, 13, 9],
                                         [9, 13, 14]])       
        
        #Neumann_BCs = np.array([[node number, in direction 0/1 , non-zero force applied (we will initialize the force vector with zeros)]]) 
        #                                     0: x1 - direction 
        #                                     1: x2 - direction 
                          
        Neumann_BCs = np.array([[1, 2, -1000],
                                 [2, 2, -2000],
                                 [3, 2, -800],
                                ])
        
        #Dirichlet_BCs = np.array([[node number, in dircection 0 or 1, prescribed displacement]])
        #                                        0: x1 - direction 
        #                                        1: x2 - direction 
        
        Dirichlet_BCs = np.array([[0, 0, 0.0],
                                [0, 1, 0.0],
                                [5, 0, 0.0],
                                [5, 1, 0.0],
                                [10, 0, 0.0],
                                [10, 1, 0.0]])

        # material properties:
        v  = 0.3  # Poisson's ratio
        t  = 0.1  # thickness
        Ey = 2E9  # Young's Modulus [N/m^2]  
    ###########################################################################################################
    elif problem == 2:    
        # you can ignore this section for now, we will get back to this later on.
        
        # Enter h, sigma1, sigma2:
        h = 0.3
        sigma1 = ...
        sigma2 = ...
        
        from lib.P4_Problem2_Input import P4_Problem2_Input
        
        nodal_positions, nodal_connectivities, Dirichlet_BCs, Neumann_BCs, Ey, v, t = \
            P4_Problem2_Input(h, sigma1, sigma2)
    ###########################################################################################################
    
    else:
        print('please enter 1 or 2!')
    return nodal_positions, nodal_connectivities, Dirichlet_BCs, Neumann_BCs, Ey, v, t
```

### select input

We define the two problems of this project below in this input section. 


:::{admonition} attention
:class: attention

Make sure to enter the problem number you are solving below. Ensure that **problem = 1 for your submission**!
:::



```python
problem = 1
nodal_positions, nodal_connectivities, Dirichlet_BCs, Neumann_BCs, Ey, v, t   = input_definition(problem)
```

```python
grader.check("input")
```

Output:
```
input results: All test cases passed!
input - 1 message: problem 1 - nice, you seem to have defined the tested input variables for problem 1 correctly.
```

### set-up

This should be familiar to you by now: you are only asked to construct the 
1. material stiffness matrix $\boldsymbol{E}_{\text{plane stress}}$ called ```E_matrix``` in the code and
2. ```elment_dof_connectivity``` matrix below: 

```python
# ----------------------------------------------------------------------------------------------------------#
#                                          Set up FE
# ----------------------------------------------------------------------------------------------------------#
no_dofs_per_node = 2                              # number of dofs per node
no_global_nodes  = nodal_positions.shape[0]       # number of global nodes
no_global_dofs   = no_global_nodes * no_dofs_per_node # number of global dofs
no_elements      = nodal_connectivities.shape[0]  # number of elements
no_Neumann_BCs   = Neumann_BCs.shape[0]           # number of external loads applied
no_Dirichlet_BCs = Dirichlet_BCs.shape[0]         # number of Dirichlet boundary conditions

# TODO: construct the material stiffness matrix plane stress:
E_matrix = Ey/(1-v**2)*np.array([[1, v, 0], 
                                [v, 1, 0], 
                                [0,0,0.5*(1-v)]])

# TODO: construct the element_dof_connectivity matrix generically (i.e. using only variable names):
nodes1 = nodal_connectivities[:, 0]
nodes2 = nodal_connectivities[:, 1]
nodes3 = nodal_connectivities[:, 2]
element_dof_connectivity = np.column_stack([
    nodes1 * no_dofs_per_node, nodes1 * no_dofs_per_node + 1, nodes1 * no_dofs_per_node + 2,
    nodes2 * no_dofs_per_node, nodes2 * no_dofs_per_node + 1, nodes2 * no_dofs_per_node + 2,
    nodes3 * no_dofs_per_node, nodes3 * no_dofs_per_node + 1, nodes3 * no_dofs_per_node + 2,
]).astype(int)
# print(element_dof_connectivity)
```

```python
grader.check("setUpFE")
```

Output:
```
setUpFE results: All test cases passed!
setUpFE - 1 message: nice, your E-matrix seems to be correct!
```

### define functions

As we have seen in Coding Lab 3, let's **write two functions** to make our life easier: fill out the missing information below to compute 
1. the $\boldsymbol{B}_e$ - matrix (per element) and
2. the element stiffness matrix $\boldsymbol{K}_e$:

```python
def compute_Be(nodal_positions, node1e, node2e, node3e):
    
    # write your code here to compute the B-matrix per element,
    # using the inputs provided: namely the nodal_positions and node1e, node2e, node3e, 
    # which are the three nodes associated with the current element. 
      
    # coordinates
    x1, y1 = nodal_positions[node1e]
    x2, y2 = nodal_positions[node2e]
    x3, y3 = nodal_positions[node3e]

    # area (signed)
    A2 = x1*(y2 - y3) + x2*(y3 - y1) + x3*(y1 - y2)   # 2*Area
    if A2 == 0:
        raise ValueError("Zero area element")
    A = 0.5 * A2

    # bi and ci coefficients (standard CST notation)
    b1 = y2 - y3
    b2 = y3 - y1
    b3 = y1 - y2
    c1 = x3 - x2
    c2 = x1 - x3
    c3 = x2 - x1

    # B matrix (3 x 6) for CST element with DOFs [u1,v1,u2,v2,u3,v3]
    Be = (1.0 / (2.0 * A)) * np.array([
        [b1,   0.0, b2,   0.0, b3,   0.0],
        [0.0,  c1,  0.0, c2,  0.0, c3 ],
        [c1,   b1, c2,   b2, c3,   b3 ]
    ])
    
    return Be
```

```python
grader.check("Be")
```

Output:
```
Be results: All test cases passed!
Be - 1 message: seems like your function to compute the Be matrix works correctly.
```

```python
def compute_element_stiffnes_matrix(nodal_positions, node1e, node2e, node3e, E_matrix, t):
    
    # write your code here to compute the elemental stiffness matrix Ke
    # using the inputs provided: namely the nodal_positions and node1e, node2e, node3e, 
    # which are the three nodes associated with the current element, 
    # as well as the E-matrix and the thickness t of the element.
    
    # coordinates
    x1, y1 = nodal_positions[node1e]
    x2, y2 = nodal_positions[node2e]
    x3, y3 = nodal_positions[node3e]

    # area (signed)
    A2 = x1*(y2 - y3) + x2*(y3 - y1) + x3*(y1 - y2)   # 2*Area
    if A2 == 0:
        raise ValueError("Zero area element")
    A = 0.5 * A2

    # compute B matrix
    Be = compute_Be(nodal_positions, node1e, node2e, node3e)
    element_stiffness_matrix = t * A * (Be.T @ E_matrix @ Be)
    
    return element_stiffness_matrix
```

```python
grader.check("Ke")
```

Output:
```
Ke results: All test cases passed!
Ke - 1 message: seems like your function to compute the Ke matrix works correctly.
```

Applying the boundary conditions and solving is given here in a function. We wrote this completely generically in coding lab 3, **so no changes are required**. You can click on the three dots to see details of the function.

```python
def apply_BCs_and_solve(global_stiffness_matrix):
    import time 
    # Convert to sparse 
    global_stiffness_matrix = csr_matrix(global_stiffness_matrix)

    # Compute DOFs
    non_zero_Neumann_BCs_dofs = (no_dofs_per_node * Neumann_BCs[:, 0].astype(int) 
                                 + Neumann_BCs[:, 1].astype(int))
    
    Dirichlet_BCs_dofs = (no_dofs_per_node * Dirichlet_BCs[:, 0].astype(int)
                          + Dirichlet_BCs[:, 1].astype(int))
    
    # Use np.setdiff1d for fast difference computation
    Neumann_BCs_dofs = list(set(range(no_global_dofs)) - set(Dirichlet_BCs_dofs))
    # Initialize global vectors
    global_force_vector        = np.zeros(no_global_dofs)
    global_displacement_vector = np.zeros(no_global_dofs)

    # Apply Neumann and Dirichlet BCs
    global_force_vector[non_zero_Neumann_BCs_dofs] = Neumann_BCs[:, 2]
    global_displacement_vector[Dirichlet_BCs_dofs] = Dirichlet_BCs[:, 2]

    # Extract submatrices
    KFF = global_stiffness_matrix[np.ix_(Neumann_BCs_dofs, Neumann_BCs_dofs)]
    KFE = global_stiffness_matrix[np.ix_(Neumann_BCs_dofs, Dirichlet_BCs_dofs)]
    
    UE = global_displacement_vector[Dirichlet_BCs_dofs]
    FF = global_force_vector[Neumann_BCs_dofs]

    # Solve the reduced system using sparse solver
    rhs = FF - KFE @ UE  # Compute this first to avoid redundant computation
    UF = spsolve(KFF, rhs)
    
    # Store results
    global_displacement_vector[Neumann_BCs_dofs] = UF

    # Solve for reaction forces
    global_force_vector = global_stiffness_matrix@global_displacement_vector

    return global_displacement_vector, global_force_vector   
```

### assemble global stiffness matrix & apply BC's and solve problem
We loop over all elements and compute the elemental stiffness matrices and assemble these to the correct positions into the global stiffness matrix, **no changes required**. Simply execute this code block.

Enforcing the essential and natural boundary conditions and computing the global displacement and force vectors.
**Note:** Since we wrote this section generically in the previous projects, we can reuse it here and **do not have to apply any changes**, regardless if we use bar or beam or CST elements. We simply make use of the function we defined earlier on. Execute the code cell below.

```python
# ----------------------------------------------------------------------------------------------------------#
#                                                ASSEMBLE K
# ----------------------------------------------------------------------------------------------------------#
#Initialize global stiffness matrix to the correct size:
global_stiffness_matrix = np.zeros((no_global_dofs, no_global_dofs))

# Assemble global stiffness matrix
# Loop over all elements -> compute elemental stiffness matrix -> assemble global stiffness matrix
for e in range(no_elements):
    node1e = nodal_connectivities[e, 0]
    node2e = nodal_connectivities[e, 1]
    node3e = nodal_connectivities[e, 2]
    element_stiffness_matrix = compute_element_stiffnes_matrix(nodal_positions, node1e, node2e, node3e, E_matrix, t)

    # Assemble Ke into Kglobal:
    global_stiffness_matrix[np.ix_(element_dof_connectivity[e], element_dof_connectivity[e])] += element_stiffness_matrix
# ----------------------------------------------------------------------------------------------------------#
#                                               APPLYING BCs & SOLVING
# ----------------------------------------------------------------------------------------------------------#
# The computation of the global displacement and force vectors does not require any changes.
global_displacement_vector, global_force_vector = apply_BCs_and_solve(global_stiffness_matrix)
# ----------------------------------------------------------------------------------------------------------#
```

Output:
```
ValueError                                Traceback (most recent call last)
Cell In[179], line 16
     13     element_stiffness_matrix = compute_element_stiffnes_matrix(nodal_positions, node1e, node2e, node3e, E_matrix, t)
     15     # Assemble Ke into Kglobal:
---> 16     global_stiffness_matrix[np.ix_(element_dof_connectivity[e], element_dof_connectivity[e])] += element_stiffness_matrix
     17 # ----------------------------------------------------------------------------------------------------------#
     18 #                                               APPLYING BCs & SOLVING
     19 # ----------------------------------------------------------------------------------------------------------#
     20 # The computation of the global displacement and force vectors does not require any changes.
     21 global_displacement_vector, global_force_vector = apply_BCs_and_solve(global_stiffness_matrix)

ValueError: operands could not be broadcast together with shapes (9,9) (6,6) (9,9)
```

## problem 1 b)
:::{admonition} submission requirement
:class: tip

Now that we solved problem 1 using CST elements, i.e. you computed all nodal displacements and forces acting on the support of the sign, we do some post-processing:


```{figure} figures/lab4_plate_stress.png
:name: roadSign
:alt: plane stress road sign
:class: bg-primary
:width: 400px
:align: center

```

Compute the **stress component** $\sigma_{11}$  in the **two elements marked in red and labelled *upper* and *lower*** in the above figure and print these to screen. 

Finally, use the provided plotting function to plot the initial and deformed structure.
:::


```python
# ----------------------------------------------------------------------------------------------------------#
#                                        POST-PROCESSING
# ----------------------------------------------------------------------------------------------------------#
if problem == 1:

    # write your code here to compute the stress component 11 in the two elements marked in red in the above figure.
    # Use the provided names sigma_11_upper and sigma_11_lower and then print these values to screen using the provided print function.
    # (as a sanity check, you can think about what the positive or negative signs you might get out of your 
    # calculation mean and whether this makes sense with regard to the applied loading condition.)
    
    elem_upper = 2
    elem_lower = 10

    # helper to compute element stress vector (sigma = E * B * ue)
    def element_stress(e):
        node_idxs = nodal_connectivities[e].astype(int)
        dof_idxs = element_dof_connectivity[e]
        ue = global_displacement_vector[dof_idxs]  # [u1,v1,u2,v2,u3,v3]
        Be = compute_Be(nodal_positions, node_idxs[0], node_idxs[1], node_idxs[2])
        sigma = E_matrix @ (Be @ ue)   # stress vector [s11, s22, s12]
        return sigma

    sigma_upper = element_stress(elem_upper)
    sigma_lower = element_stress(elem_lower)

    sigma_11_upper = sigma_upper[0]
    sigma_11_lower = sigma_lower[0]
    print(f'the stress sigma_11_upper = {sigma_11_upper:.2f} Pa')
    print(f'the stress sigma_11_lower = {sigma_11_lower:.2f} Pa')

    # simply execute this code block to see the initial and deformed structure thanks to the provided plotting function:        
    scale_factor = 100 # we scale the deformed solution so that it becomes visible seeing that we are working with small displacements           
    plot_original_and_displaced_structure(nodal_positions, global_displacement_vector, nodal_connectivities, scale_factor)
```

Output:
```
Cell In[90], line 23
     20     sigma = E_matrix @ (Be @ ue)   # stress vector [s11, s22, s12]
     21     return sigma
---> 23 sigma_upper = element_stress(elem_upper)
     24 sigma_lower = element_stress(elem_lower)
     26 sigma_11_upper = sigma_upper[0]

Cell In[90], line 18, in element_stress(e)
     16 node_idxs = nodal_connectivities[e].astype(int)
     17 dof_idxs = element_dof_connectivity[e]
---> 18 ue = global_displacement_vector[dof_idxs]  # [u1,v1,u2,v2,u3,v3]
     19 Be = compute_Be(nodal_positions, node_idxs[0], node_idxs[1], node_idxs[2])
     20 sigma = E_matrix @ (Be @ ue)   # stress vector [s11, s22, s12]

NameError: name 'global_displacement_vector' is not defined
```

___
## problem 2 a)
:::{admonition} submission requirement
:class: tip
Now we are all set to solve a real-world problem. We would like to know how well our FE code performs when it comes to stress distributions around a hole in a plate that is loaded remotely (see the figure below). The circumferential stress along the edge of a circular hole in a 2D infinite plate under remote loading was found analytically in the classical work of Savin (1961), referencing back to Muskhelishvili (1934), as

$$
	\sigma_\text{circ}(\theta)=(\sigma_1+\sigma_2)-2(\sigma_1-\sigma_2)\cos(2\theta),
$$

where $\sigma_{1,2}$ are the remotely applied stresses in the $x_{1,2}$-directions (see the figure), and the position on the circumference is indicated by $\theta$, the angle measured from the horizontal axis. 
Notice that in the special case $\sigma_{2}=0$ (uniaxial loading), the maximum stress is $\sigma_\text{circ}(\tfrac{\pi}{2})=3\sigma_1$ and $\sigma_\text{circ}(0)=-\sigma_1$.

Let us solve this boundary problem numerically with our FE code.
By symmetry, it is sufficient to simulate only one quadrant of the problem with uniform tractions applied on two edges and sliding essential boundary conditions applied on the other two edges, as shown in the figure. 

```{figure} figures/infinitePlate-eps-converted-to.png
:name: sketch
:alt: infinite plate sketch
:class: bg-primary
:width: 300px
:align: center

```

```{figure} figures/infinitePlateFE.png
:name: plateFE
:alt: infinite plate FE
:class: bg-primary
:width: 400px
:align: center

```

%<img src="figures/infinitePlate-eps-converted-to.png" alt="infinite plate sketch" style="width:300px;"/>
%<!-- <p style="font-style:italic;">Figure 3: infinite plate sketch </p> -->


%<img src="figures/infinitePlateFE.png" alt="infinite plate FE" style="width:400px;"/>
%<!-- <p style="font-style:italic;">Figure 4: infinite plate FE </p> -->

Simply specify $h,\sigma_1,\sigma_2$ in the input section and select problem = 2 to run the problem.
Young's modulus, Poisson's ratio and plate thickness are $E = 10^3$ N/m$^2$, $\nu = 0.3$ and $t = 1$ m , respectively.
Use $\sigma_1 = 10$ N/m$^2$ and $\sigma_2 = 0$ N/m$^2$.
You are aslo required to input the $h$-factor, which governs the element size around the hole.
For lower values of $h$ the mesh is more refined.
Please note that the underlying mesh refinement algorithm may produce a slightly different mesh when you run your simulation multiple times using the same value for $h$.
Run your FE code with $h = 0.08$ (you should already see some mesh refinement around the hole) and compute all element strains and stresses.
For each element, we also calculate the von Mises stress

$$
	\sigma_\text{vM} = \sqrt{\sigma_{11}^2 - \sigma_{11} \sigma_{22} + \sigma_{22}^2 + 3 \sigma_ {12}^2 },
$$

where we assume plane stress conditions ($\sigma_{33}=0$). 

The von Mises stress is an effective stress measure, which is a scalar quantity.
On the hole's circumference, the von Mises stress and the analytical solution coincide ($\sigma_{\text{circ}}=\sigma_{\text{vM}}$) since there is only one non-zero stress component.
We can verify that the maximum von Mises stress on the circumference that you compute is found in the outermost element on the circumference (at $\theta = \tfrac{\pi}{2}$), which is in agreement with the analytical solution.

---
**Input** the required $h,\sigma_1,\sigma_2$ in the pre-processing section: use $\sigma_1 = 10$ N/m$^2$ and $\sigma_2 = 0$ N/m$^2$ and  $h = 0.08$. 

Further, **select problem = 2** in the select input section.

For this $h = 0.08$, **calculate the strains and stresses for each element** and from those the **von Mises stress for each element** in the post-processing section below.
Use the provided plotting function to plot the von Mises stress throughout the domain and verify that the maximum von Mises stress indeed appears at the expected location.
:::

```python
if problem == 2:  
    # initialize a vector to store the von Mises stress per element
    # no changes required
    stress_von_Mises = np.zeros((no_elements, 1))

    # compute the von Mises stress per element and store it inside stress_von_Mises:
    for e in range(no_elements):
        # write your code here
        stress_von_Mises[e] = ...
    
    # provided plotting function
    plot_von_Mises_stress_distribution(nodal_positions,nodal_connectivities, stress_von_Mises)

    # printing the maximum stress to screen is provided here:
    print(f"Maximum von Mises stress is {np.max(stress_von_Mises):.2f} Pa")
    
```

## problem 2 b)

:::{admonition} submission requirement
:class: tip
We now want to refine the mesh around the hole. 
Run your code two more times, each time decreasing the value of $h$, using $h = 0.04$ and $h = 0.01$, respectively.
Does the FE solution improve with $h$-refinement as compared to the analytical solution?  

Use the predefined structure below to save the maximum von Mises stress (which is displayed above) corresponding to the respective level of refinement.
:::

```python
h_values                = np.array([ 0.01,0.04, 0.08 ])
# uncomment the following line of code and fill in the numerical values of the corresponding max. von Mises stress
# simply copy the value that you print to screen above.
# corr_max_von_Mises_stress = np.array([ max_sigma_vM_for_h0.01, max_sigma_vM_for_h0.04, max_sigma_vM_for_h0.08 ]) 
```

## problem 2 c)
:::{admonition} submission requirement
:class: tip
Now use the provided plotting functions to show the plot $h$ vs. $\sigma_{\text{vM max}}$.
The plot includes a line indicating the exact solution.

Simply uncomment the code cell below and run it:

:::

<!-- BEGIN QUESTION -->



```python
# simply uncomment the code and execute this code cell:
#plot_h_vs_Sigma(h_values,corr_max_von_Mises_stress)
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

## problem 2 d)

:::{admonition} submission requirement
:class: tip
Comment on the accuracy/convergence of the FE-code in the space provided below. What trend do you observe? Is this expected? Why/why not?
:::

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

***

## Closing remarks

When you reach this part of the script, your code should be running just fine.  
Further, you have taken some time to interpret some of the results. 

To close off this assignment on 2D elements, write down three points of what you have learned during this project in the provided space below.

As a final point, you can also give an indication of how long it took you to complete this assignment. 

_Type your answer here, replacing this text._

```python
grader.check("Check saved problem 1 for submission")
```

<!-- END QUESTION -->

:::{admonition} important
:class: attention

**Run** and **save** your code with **problem = 1** in code cell 4 before submitting your coding-lab-4 folder with all its dependencies through Moodle!

Best of luck!

:::

Tests for Be (points: None)
```
>>> nodal_positionstest = np.array([[2.0, 1.0], [4.0, 2.0], [3.0, 3.0]])
>>> node1etest = 0
>>> node2etest = 1
>>> node3etest = 2
>>> myBE_matrixTest = compute_Be(nodal_positionstest, node1etest, node2etest, node3etest)
>>> assert np.allclose(myBE_matrixTest, [[-0.33333333, 0.0, 0.66666667, 0.0, -0.33333333, 0.0], [0.0, -0.33333333, 0.0, -0.33333333, 0.0, 0.66666667], [-0.33333333, -0.33333333, -0.33333333, 0.66666667, 0.66666667, -0.33333333]], atol=1e-06)
```
failure_message: seems like your function to compute the Be matrix does NOT work correctly.
success_message: seems like your function to compute the Be matrix works correctly.


Tests for Check saved problem 1 for submission (points: None)
```
>>> assert problem == 1
```
failure_message: You can ignore this test for now, but make sure to run your code with problem = 1 and save before submitting!
success_message: great, you ran problem = 1 as required for submission.


Tests for Ke (points: None)
```
>>> nodal_positionstest = np.array([[2.0, 1.0], [4.0, 2.0], [3.0, 3.0]])
>>> node1etest = 0
>>> node2etest = 1
>>> node3etest = 2
>>> Eytest = 1
>>> vtest = 0.2
>>> ttest = 0.1
>>> E_matrixtest = Eytest / (1 - vtest ** 2) * np.array([[1, vtest, 0], [vtest, 1, 0], [0, 0, (1 - vtest) / 2]])
>>> ttest = 1
>>> myKE_matrixTest = compute_element_stiffnes_matrix(nodal_positionstest, node1etest, node2etest, node3etest, E_matrixtest, ttest)
>>> assert np.allclose(myKE_matrixTest, [[0.24305556, 0.10416667, -0.27777778, -0.10416667, 0.03472222, 0.0], [0.10416667, 0.24305556, 0.0, 0.03472222, -0.10416667, -0.27777778], [-0.27777778, 0.0, 0.76388889, -0.20833333, -0.48611111, 0.20833333], [-0.10416667, 0.03472222, -0.20833333, 0.45138889, 0.3125, -0.48611111], [0.03472222, -0.10416667, -0.48611111, 0.3125, 0.45138889, -0.20833333], [0.0, -0.27777778, 0.20833333, -0.48611111, -0.20833333, 0.76388889]])
```
failure_message: seems like your function to compute the Ke matrix does NOT work correctly.
success_message: seems like your function to compute the Ke matrix works correctly.


Tests for input (points: None)
```
>>> if problem == 1:
...     assert np.isclose(np.log(v) * Ey / 3672300000 + np.sqrt(nodal_positions[0, 1]) * nodal_positions[14, 0] * 1 / t, 6.415362774517799)
```
failure_message: problem 1 - please make sure that all your input variables are defined correctly. (ignore for problems 2!)
success_message: problem 1 - nice, you seem to have defined the tested input variables for problem 1 correctly.


Tests for q1 U (points: None)


Tests for q1PostProcess (points: None)


Tests for setUpFE (points: None)
```
>>> if problem == 1:
...     assert np.isclose(E_matrix.sum(), 6483516483.516484)
```
failure_message: there is a mistake in your E_matrix :-(
success_message: nice, your E-matrix seems to be correct!
